<problem linenum="23" tex_filename="./week4_5_ps2b.tex" attempts="10" display_name="Two-qubit Clifford circuits" max_attempts="10">
  <script>
from sympy_check import *

</script>
  <p>
We explore some properties of quantum circuits constructed from gates in the Clifford group. These are often called simply “Clifford circuits," for short, and here we are interested in two-qubit circuits. </p>
  <ol class="enumerate">
    <li value="1">
      <p>
Recall that the CNOT gate acts on a two-qubit Pauli group element [mathjaxinline]g\in P_2[/mathjaxinline] in a simple way, according to this table (the control qubit is on the left, and the target on the right): </p>
      <style>
.datagrid table { border-collapse: collapse; text-align: center; } .datagrid {font: normal 20px/150%; font-family:"Lucida Console", Monaco, monospace; background: #fff; overflow: hidden; border: none; -webkit-border-radius: 3px; -moz-border-radius: 3px; border-radius: 3px; }.datagrid table td, .datagrid table th { padding: 3px 10px; }.datagrid table thead th {background:-webkit-gradient( linear, left top, left bottom, color-stop(0.05, #006699), color-stop(1, #00557F) );background:-moz-linear-gradient( center top, #006699 5%, #00557F 100% );filter:progid:DXImageTransform.Microsoft.gradient(startColorstr='#006699', endColorstr='#00557F');background-color:#006699; color:#ffffff; font-size: 25px; font-weight: bold; border-left: 1px solid #0070A8; } .datagrid table thead th:first-child { border: none; }.datagrid table tbody td { color: #00496B; border-left: 1px solid #E1EEF4;font-size: 20px;font-weight: normal; }.datagrid table tbody .alt td { background: #E1EEF4; color: #00496B; }.datagrid table tbody td:first-child { border-left: none; }.datagrid table tbody tr:last-child td { border-bottom: none; }.datagrid table tfoot td div { border-top: 1px solid #006699;background: #E1EEF4;} .datagrid table tfoot td { padding: 0; font-size: 20px } .datagrid table tfoot td div{ padding: 2px; }.datagrid table tfoot td ul { margin: 0; padding:0; list-style: none; text-align: center; }.datagrid table tfoot  li { display: inline; }.datagrid table tfoot li a { text-decoration: none; display: inline-block;  padding: 2px 8px; margin: 1px;color: #FFFFFF;border: 1px solid #006699;-webkit-border-radius: 3px; -moz-border-radius: 3px; border-radius: 3px; background:-webkit-gradient( linear, left top, left bottom, color-stop(0.05, #006699), color-stop(1, #00557F) );background:-moz-linear-gradient( center top, #006699 5%, #00557F 100% );filter:progid:DXImageTransform.Microsoft.gradient(startColorstr='#006699', endColorstr='#00557F');background-color:#006699; }.datagrid table tfoot ul.active, .datagrid table tfoot ul a:hover { text-decoration: none;border-color: #006699; color: #FFFFFF; background: none; background-color:#00557F;}div.dhtmlx_window_active, div.dhx_modal_cover_dv { position: fixed !important; }
</style>
      <center style="boder:none; padding:none">
        <div class="datagrid">
          <table>
            <thead>
              <tr>
                <th>g</th>
                <th>CNOT on g</th>
              </tr>
            </thead>
            <tbody>
              <tr>
                <td>IX</td>
                <td style="text-align:center">IX</td>
              </tr>
              <tr class="alt">
                <td>IZ</td>
                <td style="text-align:center">ZZ</td>
              </tr>
              <tr>
                <td>XI</td>
                <td style="text-align:center">XX</td>
              </tr>
              <tr class="alt">
                <td>ZI</td>
                <td style="text-align:center">ZI</td>
              </tr>
            </tbody>
          </table>
        </div>
      </center>
      <p>
This table can be expressed as a map, in which we say that CNOT maps the inputs <tt class="tt">[IX,IZ,XI,ZI]</tt> to the outputs <tt class="tt">[IX,ZZ,XX,ZI]</tt>. Following this notation, give a list which represents the result of the CNOT map acting on the inputs <tt class="tt">[XY,YY,ZY]</tt>: </p>
      <p>
        <p style="display:inline"><tt class="tt">[XY,YY,ZY]</tt>[mathjaxinline]\rightarrow[/mathjaxinline]</p>
        <span>
          <customresponse cfn="cfn_wrap_b365a1273c" inline="1" expect="$expect_wrap_b365a1273c">
            <textline size="40" correct_answer="$ans_b365a1273c_1" inline="1"/>
          </customresponse>
          <script type="text/python" system_path="python_lib">

import subtext2
wrap_b365a1273c = subtext2.Subtext(debug=True,sanitize_allow_lists=False)
ans_b365a1273c_1 = wrap_b365a1273c.answer('''[YZ,-XZ,IY]''')
expect_wrap_b365a1273c = wrap_b365a1273c.answer('''[YZ,-XZ,IY]''')

@wrap_b365a1273c.grader
def cfn_wrap_b365a1273c(expect, ans, **kwargs):
    return sympy_formula_check(expect, ans, **kwargs)

</script>
        </span>
      </p>
    </li>
    <li value="2">
      <p>
Consider the circuit <tt class="tt">[ CNOT(0,1), CNOT(1,0), CNOT(0,1) ]</tt>: </p>
      <p>
        <center>
          <img src="/static/images/threecnot.png" width="300"/>
        </center>
      </p>
      <p>
Give the list which represents the output of this Clifford circuit acting on the inputs <tt class="tt">[IX,IZ,XI,ZI]</tt>: </p>
      <p>
        <p style="display:inline"><tt class="tt">[IX,IZ,XI,ZI]</tt>[mathjaxinline]\rightarrow[/mathjaxinline]</p>
        <span>
          <customresponse cfn="cfn_wrap_65953b0b93" inline="1" expect="$expect_wrap_65953b0b93">
            <textline size="40" correct_answer="$ans_65953b0b93_1" inline="1"/>
          </customresponse>
          <script type="text/python" system_path="python_lib">

import subtext2
wrap_65953b0b93 = subtext2.Subtext(debug=True,sanitize_allow_lists=False)
ans_65953b0b93_1 = wrap_65953b0b93.answer('''[XI,ZI,IX,IZ]''')
expect_wrap_65953b0b93 = wrap_65953b0b93.answer('''[XI,ZI,IX,IZ]''')

@wrap_65953b0b93.grader
def cfn_wrap_65953b0b93(expect, ans, **kwargs):
    return sympy_formula_check(expect, ans, **kwargs)

</script>
        </span>
      </p>
    </li>
  </ol>
</problem>
