<problem linenum="6" tex_filename="./week1_3_denmat.tex" attempts="10" display_name="Density matrices and unravelings (1 of 3)" max_attempts="10">
  <p>
This problem illustrates the many possible physical interpretations of a density matrix, based on the infinite number possible <em>unravelings</em> of density matrices into different possible statistical mixtures of pure states. </p>
  <p>
Consider the bipartite quantum state </p>
  <table id="a0000000013" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
    <tr>
      <td class="equation" style="width:80%; border:none">[mathjax]|\psi _{AB}\rangle  = \sqrt{\frac{3}{4}} |00\rangle  + \sqrt{\frac{1}{4}} |11\rangle[/mathjax]</td>
      <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.1)</td>
    </tr>
  </table>
  <p>
jointly held by Alice and Bob (each has one qubit). Suppose Bob measures his qubit in the computational basis. We may depict this with the quantum circuit: </p>
  <p>
    <center>
      <img src="/static/images/denmat-measure-basis1.png" width="420"/>
    </center>
  </p>
  <p>
What state does Alice have? Answer this by writing down, at first, Alice's state [mathjaxinline]|\psi '_{A,k}\rangle[/mathjaxinline] (after Bob's measurement) conditioned on Bob obtaining result [mathjaxinline]k[/mathjaxinline] (either [mathjaxinline]k=0[/mathjaxinline] or [mathjaxinline]k=1[/mathjaxinline]) from his measurement; include the probability of this result occuring. </p>
  <p>
Be sure that [mathjaxinline]|\psi '_{A,k}\rangle[/mathjaxinline] is properly normalized. </p>
  <p>
Recall that states should be entered using “ket" notation, e.g. [mathjaxinline]|0\rangle[/mathjaxinline] is [mathjaxinline]{\tt |0&gt;}[/mathjaxinline]. </p>
  <script>
from sympy_check import *

</script>
  <script linenum="45" filename="./week1_3_denmat.tex" type="text/python" system_path="python_lib">

myhints = [ {'eval': "not symbol('g')", 'hint': "Shouldn't your answer depend on g?"},
            {'eval': "not string('sqrt')", 'hint': "What should the answer be for g=0 and g=1?"},
          ]
  
</script>
  <ul class="itemize">
    <li>
      <p>
        <p style="display:inline">Alice's state [mathjaxinline]|\psi '_{A,0}\rangle[/mathjaxinline] when Bob measures [mathjaxinline]k=0[/mathjaxinline]:</p>
        <customresponse cfn="sympy_check_quantum" inline="1" expect="|0&amp;gt;" options="phase_ambiguity" cfn_extra_args="options">
          <textline size="60" correct_answer="|0&amp;gt;" inline="1" math="1" preprocessorClassName="MathjaxPreprocessorForQM" preprocessorSrc="/static/mathjax_preprocessor_for_QM_H.js"/>
        </customresponse>
      </p>
    </li>
    <li>
      <p>
        <p style="display:inline">Probability of Bob measuring [mathjaxinline]k=0[/mathjaxinline]:</p>
        <customresponse cfn="sympy_formula_check" inline="1" expect="3/4">
          <textline size="10" correct_answer="3/4" inline="1"/>
        </customresponse>
      </p>
    </li>
    <li>
      <p>
        <p style="display:inline">Alice's state [mathjaxinline]|\psi '_{A,1}\rangle[/mathjaxinline] when Bob measures [mathjaxinline]k=1[/mathjaxinline]:</p>
        <customresponse cfn="sympy_check_quantum" inline="1" expect="|1&amp;gt;" options="phase_ambiguity" cfn_extra_args="options">
          <textline size="60" correct_answer="|1&amp;gt;" inline="1" math="1" preprocessorClassName="MathjaxPreprocessorForQM" preprocessorSrc="/static/mathjax_preprocessor_for_QM_H.js"/>
        </customresponse>
      </p>
    </li>
    <li>
      <p>
        <p style="display:inline">Probability of Bob measuring [mathjaxinline]k=1[/mathjaxinline]:</p>
        <customresponse cfn="sympy_formula_check" inline="1" expect="1/4">
          <textline size="10" correct_answer="1/4" inline="1"/>
        </customresponse>
      </p>
    </li>
  </ul>
</problem>
