<problem linenum="625" tex_filename="./week3_4_ps2a.tex" attempts="10" display_name="Four-qubit stabilizer state II" max_attempts="10">
  <script>
import qpython4
qpython4.define_qpython4_globals(globals())


</script>
  <p>
Give two four-qubit states which are stabilized by [mathjaxinline]{\cal S}[/mathjaxinline]. Please input your answer using “ket" notation, eg <tt class="tt">(|0011&gt;+|0000&gt;)/sqrt(2)</tt>. The two states should be orthogonal to each other. They need not be properly normalized. </p>
  <p>
    <customresponse cfn="check_stabilizer_state_pair" inline="1" expect="See solutions" options="stab=[IIZZ,ZZII,XXXX]" cfn_extra_args="options">
      <p style="display:inline">[mathjaxinline]|0_ L\rangle  =[/mathjaxinline] <textline size="60" correct_answer="See soltuions" inline="1" math="1" preprocessorClassName="MathjaxPreprocessorForQM" preprocessorSrc="/static/mathjax_preprocessor_for_QM_H.js"/></p>
      <br/>
      <p style="display:inline">[mathjaxinline]|1_ L\rangle  =[/mathjaxinline] <textline size="60" correct_answer="See solutions" inline="1" math="1" preprocessorClassName="MathjaxPreprocessorForQM" preprocessorSrc="/static/mathjax_preprocessor_for_QM_H.js"/></p>
    </customresponse>
  </p>
  <p>
    <solution>
      <p>
        <b>Solution:</b>
      </p>
      <span>
        <table id="a0000000012" class="eqnarray" cellspacing="0" cellpadding="7" width="100%" style="table-layout:auto">
          <tr id="a0000000013">
            <td style="width:40%; border:none"> </td>
            <td style="vertical-align:middle;                                    text-align:right; border:none">
	    [mathjaxinline]\displaystyle |0_ L\rangle[/mathjaxinline]
        </td>
            <td style="vertical-align:middle;                                    text-align:center; border:none">
	    [mathjaxinline]\displaystyle =[/mathjaxinline]
        </td>
            <td style="vertical-align:middle;                                    text-align:left; border:none">
	    [mathjaxinline]\displaystyle  \frac{|0000\rangle  + |1111\rangle }{\sqrt{2}}[/mathjaxinline]
        </td>
            <td style="width:40%; border:none"> </td>
            <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.9)</td>
          </tr>
          <tr id="a0000000014">
            <td style="width:40%; border:none"> </td>
            <td style="vertical-align:middle;                                    text-align:right; border:none">
	    [mathjaxinline]\displaystyle |1_ L\rangle[/mathjaxinline]
        </td>
            <td style="vertical-align:middle;                                    text-align:center; border:none">
	    [mathjaxinline]\displaystyle =[/mathjaxinline]
        </td>
            <td style="vertical-align:middle;                                    text-align:left; border:none">
	    [mathjaxinline]\displaystyle  \frac{|0011\rangle  + |1100\rangle }{\sqrt{2}}[/mathjaxinline]
        </td>
            <td style="width:40%; border:none"> </td>
            <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.10)</td>
          </tr>
        </table>
      </span>
    </solution>
  </p>
</problem>
