<problem linenum="168" tex_filename="./week1_3_denmat.tex" attempts="10" display_name="Density matrices and unravelings (3 of 3)" max_attempts="10">
  <p>
The two sets of states [mathjaxinline]\{ |\psi '_{A,k}\rangle \}[/mathjaxinline] and [mathjaxinline]\{ |\psi "_{A,k}\rangle \}[/mathjaxinline] look quite different, somehow implying that the state Alice holds after Bob's measurement is different, and depends on what Bob's measurement result is, and also how Bob chooses to measure his qubit. </p>
  <p>
However, it is conceptually problematic to think that Alice's state depends in any way on Bob's measurement. What if Bob chose how to do his measurement long after Alice inquires about the state of her qubit? What if Alice and Bob were far apart – so fat away from each other that they could each do operations on their own qubits faster than information (eg carried at the speed of light) could travel between them? It should not be possible for Bob to be able to signal Alice faster than the speed of light! </p>
  <p>
This apparant paradox can be manifestly resolved (or more accurtately, perhaps, <em>avoided</em>) by using a better mathematical tool to represent the <em>statistical mixture</em> which describes each of the ensemble of states Alice may hold, in the two scenarios above. For an ensemble of states [mathjaxinline]|\psi _ k\rangle[/mathjaxinline] which occur with probability [mathjaxinline]p_ k[/mathjaxinline], the density matrix is </p>
  <table id="a0000000014" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
    <tr>
      <td class="equation" style="width:80%; border:none">[mathjax]\rho = \sum p_ k |\psi _ k\rangle \langle \psi _ k|[/mathjax]</td>
      <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.2)</td>
    </tr>
  </table>
  <p>
where [mathjaxinline]|\psi _ k\rangle \langle \psi _ k|[/mathjaxinline] is an <a href="http://en.wikipedia.org/wiki/Outer_product" target="_blank"><em>outer product</em></a> of [mathjaxinline]|\psi _ k\rangle[/mathjaxinline] and [mathjaxinline]\langle \psi _ k|[/mathjaxinline]. For example, </p>
  <table id="a0000000015" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
    <tr>
      <td class="equation" style="width:80%; border:none">[mathjax]|0\rangle \langle 0| = \left[ \begin{array}{cc}{1}&amp; {0}\\ {0}&amp; {0}\end{array}\right][/mathjax]</td>
      <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.3)</td>
    </tr>
  </table>
  <p>
and </p>
  <table id="a0000000016" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
    <tr>
      <td class="equation" style="width:80%; border:none">[mathjax]|0\rangle \langle 1| = \left[ \begin{array}{cc}{0}&amp; {1}\\ {0}&amp; {0}\end{array}\right][/mathjax]</td>
      <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.4)</td>
    </tr>
  </table>
  <p>
Think of the column labels as the inputs, and the row labels as the outputs. </p>
  <p>
Give the density matrices for the ensembles resulting in the two scenarios above. Recall that matrices are specified using an input which is a list of row vectors, eg [mathjaxinline]|0\rangle \langle 1|[/mathjaxinline] is <tt class="tt">[[0,1],[0,0]]</tt>: </p>
  <script>
from sympy_check import *

</script>
  <script linenum="221" filename="./week1_3_denmat.tex" type="text/python" system_path="python_lib">

myhints = [ {'eval': "not symbol('g')", 'hint': "Shouldn't your answer depend on g?"},
            {'eval': "not string('sqrt')", 'hint': "What should the answer be for g=0 and g=1?"},
          ]
  
</script>
  <ul class="itemize">
    <li>
      <p>
        <p style="display:inline">Density matrix [mathjaxinline]\rho '[/mathjaxinline] for Alice's state [mathjaxinline]|\psi '_{A,k}\rangle[/mathjaxinline]:</p>
        <customresponse cfn="sympy_formula_check" inline="1" expect="[[3/4,0],[0,1/4]]">
          <textline size="60" correct_answer="[[3/4,0],[0,1/4]]" inline="1" math="1"/>
        </customresponse>
      </p>
    </li>
    <li>
      <p>
        <p style="display:inline">Density matrix [mathjaxinline]\rho "[/mathjaxinline] for Alice's state [mathjaxinline]|\psi "_{A,k}\rangle[/mathjaxinline]:</p>
        <customresponse cfn="sympy_formula_check" inline="1" expect="[[3/4,0],[0,1/4]]">
          <textline size="60" correct_answer="[[3/4,0],[0,1/4]]" inline="1" math="1"/>
        </customresponse>
      </p>
    </li>
  </ul>
  <p>
    <solution>
      <p>
        <b>Solution:</b>
      </p>
      <span>
        <p>
Note that [mathjaxinline]\rho '[/mathjaxinline] is identical to [mathjaxinline]\rho "[/mathjaxinline]. The two states are the same! In fact, there is in principle <em>no possible physical experiment</em> which could be performed by Alice acting alone (with no information from Bob) to distinguish the two ensembles of states she holds in the two scenarios above. Her state, as described by a density matrix, is <em>independent</em> of Bob's measurement! </p>
        <p>
More generally, [mathjaxinline]\rho = \sum _ k p_ k |\psi _ k\rangle \langle \psi _ k|[/mathjaxinline] is known as an <em>unraveling</em> of [mathjaxinline]\rho[/mathjaxinline] into a statistical mixture of pure states. This is not unique, however, because in general there are an <em>infinite</em> number of possible unravelings of a given density matrix. </p>
        <p>
For example, [mathjaxinline]\rho = \rho ' = \sum _ k q_ k |\phi _ k\rangle \langle \phi _ k|[/mathjaxinline] if </p>
        <table id="a0000000017" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
          <tr>
            <td class="equation" style="width:80%; border:none">[mathjax]\sqrt{p_ k}|\psi _ k\rangle  = \sum _ j u_{kj}\sqrt{q_ j}|\phi _ j\rangle  \, ,[/mathjax]</td>
            <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.5)</td>
          </tr>
        </table>
        <p>
where [mathjaxinline]u_{kj}[/mathjaxinline] are elements of a unitary matrix. </p>
      </span>
    </solution>
  </p>
</problem>
