<problem linenum="433" tex_filename="./week4_5_ps2b.tex" attempts="10" display_name="Decomposition exercise: 2-qubit Clifford circuit" max_attempts="10">
  <script>
import qpython4
qpython4.define_qpython4_globals(globals())


</script>
  <p>
[mathjaxinline]U[/mathjaxinline] is a two-qubit Clifford circuit which maps these Pauli operators: </p>
  <table id="a0000000013" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
    <tr>
      <td class="equation" style="width:80%; border:none">[mathjax]{\tt [IX,IZ,XI,ZI]} \longrightarrow {\tt [-YY,ZI,ZZ,ZX]} \, .[/mathjax]</td>
      <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.10)</td>
    </tr>
  </table>
  <p>
Let us systematically construct a circuit for [mathjaxinline]U[/mathjaxinline] using CNOT, SWAP, [mathjaxinline]H[/mathjaxinline], and [mathjaxinline]S[/mathjaxinline], using the procedure from <a href="/course/jump_to_id/s12-wk6-clifford-decomposition" target="_blank">the previous problem</a>. </p>
  <p>
Start with qubit [mathjaxinline]0[/mathjaxinline] (the right-most label). Give a Clifford circuit [mathjaxinline]W_1[/mathjaxinline], using just SWAP, [mathjaxinline]H[/mathjaxinline], and [mathjaxinline]S[/mathjaxinline] gates, to bring the set of Pauli operators into a form such that for [mathjaxinline]U_1 = W_1 U[/mathjaxinline], </p>
  <table id="a0000000014" class="eqnarray" cellspacing="0" cellpadding="7" width="100%" style="table-layout:auto">
    <tr id="a0000000015">
      <td style="width:40%; border:none"> </td>
      <td style="vertical-align:middle;                                    text-align:right; border:none">
	    [mathjaxinline]\displaystyle U_1 {\tt IX}\,  U_1^\dagger[/mathjaxinline]
        </td>
      <td style="vertical-align:middle;                                    text-align:center; border:none">
	    [mathjaxinline]\displaystyle =[/mathjaxinline]
        </td>
      <td style="vertical-align:middle;                                    text-align:left; border:none">
	    [mathjaxinline]\displaystyle  h \otimes {\tt Z}[/mathjaxinline]
        </td>
      <td style="width:40%; border:none"> </td>
      <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.11)</td>
    </tr>
    <tr id="a0000000016">
      <td style="width:40%; border:none"> </td>
      <td style="vertical-align:middle;                                    text-align:right; border:none">
	    [mathjaxinline]\displaystyle U_1 {\tt IZ}\,  U_1^\dagger[/mathjaxinline]
        </td>
      <td style="vertical-align:middle;                                    text-align:center; border:none">
	    [mathjaxinline]\displaystyle =[/mathjaxinline]
        </td>
      <td style="vertical-align:middle;                                    text-align:left; border:none">
	    [mathjaxinline]\displaystyle  g \otimes {\tt X} \, .[/mathjaxinline]
        </td>
      <td style="width:40%; border:none"> </td>
      <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.12)</td>
    </tr>
  </table>
  <p>
where [mathjaxinline]h[/mathjaxinline] and [mathjaxinline]g[/mathjaxinline] are single qubit Pauli operators acting on qubit [mathjaxinline]1[/mathjaxinline]. </p>
  <p>
Enter your answer as a list of gates, in the form <big class="xlarge"><tt class="tt">[H(0),SWAP(0,1)]</tt></big>: </p>
  <p>
    <p style="display:inline">[mathjaxinline]W_1 =[/mathjaxinline]</p>
    <span>
      <customresponse cfn="cfn_wrap_6f1874be83" inline="1" expect="$expect_wrap_6f1874be83" options="nqubits=2;pset='-YY,[-i]*.Z:ZI,[-i]*.X'" cfn_extra_args="options">
        <textline size="50" correct_answer="$ans_6f1874be83_1" inline="1"/>
      </customresponse>
      <script type="text/python" system_path="python_lib">

import subtext2
wrap_6f1874be83 = subtext2.Subtext(debug=True,sanitize_allow_lists=False)
ans_6f1874be83_1 = wrap_6f1874be83.answer('''See solution''')
expect_wrap_6f1874be83 = wrap_6f1874be83.answer('''See solution''')

@wrap_6f1874be83.grader
def cfn_wrap_6f1874be83(expect, ans, **kwargs):
    return check_clifford_paulis(expect, ans, **kwargs)

</script>
    </span>
  </p>
  <p>
    <solution>
      <p>
        <b>Solution:</b>
      </p>
      <span>
        <p>
An acceptable answer is <tt class="tt">[SWAP(1,0),S(0), H(0)]</tt>. </p>
        <p>
This maps: </p>
        <table id="a0000000017" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
          <tr>
            <td class="equation" style="width:80%; border:none">[mathjax]{\tt [-YY, ZI, ZZ, ZX]} \longrightarrow {\tt [YZ, IX, ZX, XX]} \, .[/mathjax]</td>
            <td class="eqnnum" style="width:20%; border:none;text-align:right">(1.13)</td>
          </tr>
        </table>
      </span>
    </solution>
  </p>
</problem>
